// This shopping list program isn't compiling!
// Use your knowledge of generics to fix it.

<<<<<<< HEAD
=======
// Execute `rustlings hint generics1` for hints!

// I AM NOT DONE
>>>>>>> 4.6.0

fn main() {
    let mut shopping_list: Vec<String> = Vec::new();
    shopping_list.push("milk".into());
}
